---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Insert exercise title here

```yaml
type: VideoExercise
key: f0b83a3301
lang: r
xp: 50
skills: 1
video_link: //player.vimeo.com/video/298272711
```


---

## An exercise

```yaml
type: NormalExercise
key: c2966d938a
xp: 100
```

Please do

`@instructions`
Do this

`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}
2 + 2
```

`@solution`
```{r}

```

`@sct`
```{r}

```
